```python
# Initialize Otter
import otter
grader = otter.Notebook("HW3.ipynb")
```

# Homework 3

This notebook contains Python problems you need to complete. Write your codes directly here. Write your codes directly here. Save your code and upload this file ("HW3.ipynb") to [Gradescope](https://www.gradescope.com). 

Note that there are some hidden tests, so you may not get full points if these tests fail. Do not be misled by the message "q# passed!".

### This assignment must be completed individually.
- Review the syllabus to avoid any academic integrity issues. 
- All codes in this Jupyter notebook file should be your own creation. 
- You are encouraged to openly collaborate on **Piazza** with your peers. Therefore, any code snippets shared on Piazza will not be subject to plagiarism even if code similarity checking tool discovers identical codes or similar coding patterns.
- You may lose 2 points for each error that prevents Gradescope from auto-grading your submission.  It would be wise to comment out the lines that cause errors if you cannot fix them.

### Common errors and how to fix/prevent them
The following errors seem to be common among students. So, tips on how to fix or prevent them are given here.
- typos:  Take advantage of auto completion using the \[TAB\] key.
- incorrect use of symbols: Be careful match the opening and closing quotations, parentheses, brackets, braces. Using too many parentheses often create syntax errors.  Create some temporary variables may reduce such errors.
- wrong indentation: the error message may provide where it occured.
- incorrect data type: Use the type conversion functions, such as int(), float(), str(). I have seem many people not wrapping their name and id with quotation marks for the pledge part.
- other syntax errors: Read the error message and take a look at where it occured.  Most of the time, error messages help you fix the error.
- logic errors: Even when no syntax error is found, your answer can be marked wrong if you have logic errors. Sometimes, you can get hints from grader.check() output. The best way to fix would be inserting testing codes such as printing object, calling the function you create with sample inputs. Fixing logic errors is a challenge for even professional coders.
- submitting a wrong file to Gradescope: Make sure you follow the instructions carefully. Know how to browse to a file.  Also, try not to create multiple files and check the time stamp if you do.

```python
%pprint # to turn off pretty printing
```

**Question 1 (2pts):** 

Create a list named `four_to_eighty` that contains numbers starting from 4, ending
with 80, and increasing in multiples of 4. You may use range() function, but make sure that the type is `list`. 

```python
# Complete your code below. 
four_to_eighty = list(range(4,81,4))
```

```python
# Check if you get the correct list
four_to_eighty
```

```python
grader.check("q1")
```

**Question 2 (2pts):** 

From `four_to_eighty` above, select only the numbers divisible by 6 and assign them to a
new list named `div_by_6`. You may want to use a `for` loop with the list `four_to_eighty` as shown below. Inside the loop, make use of % and == operators to test the divisibility of each element in the list and then append the number to the list.

Start with the code below

```
# An empty list named div_by_6
div_by_6 = list()
for number in four_to_eighty:
```

```python
# Write your code below. 
# Ellipsis (...) is meant to be removed. It can generate syntax errors if you don't remove it.
div_by_6 = list()
for number in four_to_eighty:
    if (number%6) == 0:
        div_by_6.append(number)
```

```python
# Check your answer
div_by_6
```

```python
grader.check("q2")
```

**Question 3 (3pts):** 

Write a function named `get_div3or4` that takes a list of integers and returns a list of elements that are divisible by three or four.  The header of the function is specified below. 

```
# Start with the function below
def get_div3or4(inList):
    outList = list()
```

```python
# Write your code below. 
def get_div3or4(inList):
    outList = list()
    for number in inList:
        if (number%3==0) or (number%4==0):
            outList.append(number)
    return outList
```

```python
# Test if it works
get_div3or4([10,11,12,13,14,15,16])
```

```python
grader.check("q3")
```

**Question 4 (3pts):** 

Write your own median function named `my_median` for a list.  For a list with odd number of elements, the median is $x_{median}$ = $x_{(\frac{n+1}{2})}$ where $x_{(1)}, x_{(2)}, \ldots, x_{(n)}$ are the elements sorted from the least to the greatest. Please note the difference in starting index when you work with indices in Python. 


A list with even number of elements should take the average of two numbers in the middle. 
$$
x_{median} = \frac{x_{(\frac{n}{2})}+x_{(\frac{n}{2}+1)}}{2}
$$

This should be a full implementation of `my_median`. Hence, make sure that it works with lists with both odd and even number of elments. You may assume that there is at least one number in the list. 

(Optional for 1 extra credit point) In addition, if there are elements that are not comparable (e.g., complex number or None), ignore them and return the median of numbers only. 
In order to handle incomparable numbers, you may want to use a for loop to filter out numbers of int or float type and store in a new array. 

Start with this code:

```
def my_median(inputList):
```

```python
# Write your code below. 
def my_median(inputList):
    cleanList = [i for i in inputList if type(i)==int or type(i)==float]
    cleanList.sort()
    length = len(cleanList)
    if(len(cleanList)%2==0):
        median=(cleanList[int(length/2-1)]+cleanList[int(length/2)])/2
    else:
        median = cleanList[int((length+1)/2)-1]
    return(median)
```

```python
# Display and check if your answer works.
print(my_median([5, 2, 6.9, 3, 4, 0.1]), my_median([0,1,2,3,4]))
```

```python
grader.check("q4")
```

**Question 5 (5pt):** 

Write a function named `hex2dec()` that takes a hexadecimal number as a string (e.g., 'FF01A5', 'fa3b05') and converts it to a decimal number as integer.

One way of doing this would be :
1. Create a dictionary called `hexDict` with letters "0" to "F" as keys and 0 to 15 as values. (already coded for you)
2. Reverse the string and make it uppercase letters
3. Initialize a variable (e.g., decNum) for a resulting decimal number to 0. 
4. Use a for loop with an index going from 0 to (length of the string)-1 (i.e., for index in range(len(hexStr)): ).  Within the loop, keep adding the $(hex number) * 16^{(index)}$ to the decimal number.
5. Return the final decimal number

If you are not sure about the logic watch the following video
https://www.youtube.com/watch?v=pg-HEGBpCQk

Copy & paste the following code to begin your solution:

```
def hex2dec(hexString):
# A dictionary to convert hexadecimal character to integer is provided below
# \ can be used to contune your code to the next line, but it is not necessary in an IPython cell.
    hexdict = {'0':0, '1':1, '2':2, '3':3, '4':4, \
               '5':5, '6':6, '7':7,'8':8, '9':9,\
               'A':10, 'B':11, 'C':12, 'D':13, 'E':14, 'F':15}
# Complete the function.  You may not need more than 5 lines.
```

```python
# Write your code below. 
def hex2dec(hexString):
    hexdict = {'0':0, '1':1, '2':2, '3':3, '4':4, \
               '5':5, '6':6, '7':7,'8':8, '9':9,\
               'A':10, 'B':11, 'C':12, 'D':13, 'E':14, 'F':15}
    hexString=hexString[::-1].upper()
    decNum=0
    for i in range(len(hexString)):
        decNum+=hexdict[hexString[i]]*(16**i)
    return decNum
```

```python
# See if it works.
print(hex2dec('12'), hex2dec('ff'), hex2dec('A0'))
```

```python
grader.check("q5")
```

**Question 6 (4pts):** 

Write a function that checks if an integer is a prime number or not.  The function will return `True` if the integer is a prime, `False` otherwise.  If the input is not an integer greater than 1, the function displays a message, "___ is not an integer greater than 1." and returns nothing. 

*(Optional)* Enhance the function so that the function still works if an integer is passed as a string.  Any other inputs will display the same message, "___ is not an integer greater than 1." and returns nothing.

Start with the following code:

```
def is_prime(n):
    # Check if a number is a float or 1.
    if (type(n)==float or n==1):
        print(str(n) + " is not an integer greater than 1.")
        return
```

```python
# Write your code below. 
def is_prime(n):
    if (type(n)==float or n==1):
        print(str(n) + " is not an integer greater than 1.")
        return


    try:
        n=int(n)
    except:
        print(str(n) + " is not an integer greater than 1.")
        return

    result=True
    for i in range(2,n):
        if(n%i==0):
            result=False
            break

    return result
```

```python
# See if correct boolean value or None is returned.

print(is_prime(3.3))
print(is_prime(17))
print(is_prime(1))
```

```python
grader.check("q6")
```

**Question 7 (3pts):** 

- Create a numpy array of 40 random numbers drawn from a normal distribution with a mean of $\mu = 100$ and a standard deviation $\sigma = 10$. And then, assign it to a variable `random_sample`.  Hint: Use the numpy package's [random.normal()](https://numpy.org/doc/stable/reference/random/generated/numpy.random.normal.html) function.

- From `random_sample`, create another numpy array  (i.e., array with values `True` or `False` only). The value should be `True` if the corresponding number in `random_sample` is greater than or equal to the mean (i.e., 100). Assign the logical array to a variable `is_ge_mean`. 

- From `random_sample`, create another numpy array where the values are less than the mean (i.e., 100). Assign it to a variable named `low_sample`.

```python
# Import numpy package
import numpy as np
```

```python
# Write your code below. 
random_sample = np.random.normal(100,10,40)
is_ge_mean = random_sample >= 100
low_sample = random_sample[random_sample <100]
```

```python
# Check your work
print(random_sample)
print(is_ge_mean)
print(low_sample)
```

```python
grader.check("q7")
```

**Question 8 (3pts):** 

Compute the following statistics based on the given formula and store them in variables named `sample_mean`, `sample_var`, and `sample_sd`. Note that $x_1, x_2,\ldots, x_N$ are the 25 random numbers you have drawn from the normal distribution and $N-1$ is used in the denominator for the calculation of sample variance.

- sample mean, $\bar{x} = \frac{\sum_{i=1}^{N} x_i}{N}$
- sample variance, $s^2 = \frac{1}{N-1}\sum_{i=1}^{N} (x_i-\bar{x})^2$
- sample standard deviation, $s= \sqrt{s^2}$

Compare the mean and standard deviation computed above with the mean and standard deviation of the original normal distribution. See how close these values are to those of the original distribution (i.e., population mean  of 100 and population standard deviation of 10.) You may try increasing the sample size to a large number (e.g., 100,000 instead of 25) if they do not look similar to each other. But, do not include such coding in your submission.	

```python
# Write your code below.
sample_mean = random_sample.sum()/len(random_sample)
sample_var = np.sum((random_sample - sample_mean)**2)/(len(random_sample)-1)
sample_sd = np.sqrt(sample_var)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-46400288cb6a> in <module>
      1 # Write your code below.
----> 2 sample_mean = random_sample.sum()/len(random_sample)
      3 sample_var = np.sum((random_sample - sample_mean)**2)/(len(random_sample))
      4 sample_sd = np.sqrt(sample_var)
      5 print(sample_mean)

NameError: name 'random_sample' is not defined
```

```python
# Check if the values are reasonable
sample_mean, sample_var, sample_sd
```

```python
grader.check("q8")
```

**Question 9 (5pts):** 

### Working with two-dimensional numpy arrays
- Create a matrix (two-dimentional numpy array) of even numbers from 2 through 60 with 6 rows and 5 columns. The values within the matrix should increase row-wise. Assign it to a variable `even_matrix`.
- Create another matrix of odd numbers from 1 through 59 with 6 rows and 5 columns. The values within the matrix should increase row-wise . Assign it to a variable `odd_matrix`.
- Subtract odd_matrix from even_matrix, and store it in a variable named `diff_mat`. Calculate the sum of all elements in `diff_mat` and store it to `sum_of_differences`.
- Combine the rows of odd_matrix with that of even_matrix (i.e., the resulting matrix will have 12 rows and 5 columns), showing odd_matrix on top. Assign the result to a variable `combined_by_row`. Hint: Use numpy's `append()` function and specify the axis argument.
- Combine the columns of odd_matrix with that of even_matrix (i.e., the resulting matrix will have 6 rows and 10 columns), showing odd_matrix on the left. Assign the result to a variable `combined_by_col`.

```python
# Write your code below. 
even_matrix = np.array(range(2,61,2)).reshape(6,5)
odd_matrix = np.array(range(1,60,2)).reshape(6,5)
diff_mat = even_matrix-odd_matrix
sum_of_differences=diff_mat.sum()
combined_by_row = np.append(odd_matrix,even_matrix, axis=0)
combined_by_col = np.append(odd_matrix,even_matrix, axis=1)
```

```python
# Check if you have the correct matricies.
print(even_matrix)
print(odd_matrix)
print(diff_mat)
print(combined_by_row)
print(combined_by_col)
```

```python
grader.check("q9")
```

**Question 10 (5pts):** 

Do the following problems utilizing the numpy package:

- Generate an array of 15 random numbers from from 0 (inclusive) to 5 (exclusive) using numpy's [random.randint()](https://numpy.org/doc/stable/reference/random/generated/numpy.random.RandomState.randint.html#numpy.random.RandomState.randint) function. Assign it to a variable `rand_array1`. Hint: Look up the documentation to figure out how many and what arguments you need to pass to the randint() function.
- Generate another array of 15 random numbers from 5 (inclusive) to 10 (exclusive)  and assign it to a variable `rand_array2`.
- Create a matrix (two-dimentional numpy array) of 5 rows and 3 columns using the first array and assign it to a variable named `mat1`.  Create another matrix of 3 rows and 5 columns using the second array and assign it to a varaible named `mat2`.
- Calculate the sum of the products of the elements in the $3^{rd}$ row of `mat1` and the corresponding elements in the $4^{th}$ column of `mat2`. Note that the $3^{rd}$ row should be extracted with `mat1[2,:]` instead of `mat1[3,:]` because index starts from 0 in Python. Similarly, you need to use `mat2[:,3]` to extract the $4^{th}$ column. Assign the calulated value to a variable named `a34`.

```python
# Write your code below. 
rand_array1 = np.random.randint(0,5,15)
rand_array2 = np.random.randint(5,10,15)
mat1 = rand_array1.reshape(5,3)
mat2 = rand_array2.reshape(3,5)
a34 = np.sum(mat1[2,:]*mat2[:,3])
```

```python
# Display and check if your answer works.
print(mat1)
print(mat2)
print(a34)
print(np.matmul(mat1,mat2))
```

```python
grader.check("q10")
```

**Honor Pledge:** As a signature for the following honor pledge, store your name as it is appeared on Canvas (e.g., "Mike Jordan") and Login ID (e.g., "mjorda21") in variables `my_name` and `my_id` below.  

***I affirm that I have not given or received any unauthorized help on this assignment and that this work is my own.***

```python
my_name = "Faye Struble"
my_id = "fstruble"
```

Tests for pledge (points: [0.5])


Tests for q1 (points: [1, 1])
```
>>> type(four_to_eighty)==list and len(four_to_eighty)==20
True
```
```
>>> four_to_eighty==list(range(4,81,4))
True
```


Tests for q10 (points: [1, 1, 0.5, 0.5, 1, 0.5, 0.5])
```
>>> all((rand_array1>=0) & (rand_array1<=4))
True
```
```
>>> all((rand_array2>=5) & (rand_array2<=9))
True
```
```
>>> mat1.shape==(5,3)
True
```
```
>>> mat2.shape==(3,5)
True
```
```
>>> a34==np.matmul(mat1,mat2)[2,3]
True
```
```
>>> (mat1 == rand_array1.reshape(5,3)).all()
True
```
```
>>> np.all(mat2 == rand_array2.reshape(3,5))
True
```


Tests for q2 (points: [1, 1])
```
>>> div_by_6[len(div_by_6)-1]==72
True
```
```
>>> len(div_by_6)==6
True
```


Tests for q3 (points: 3)
```
>>> get_div3or4([10,11,12,13,14,15,16])==[12, 15, 16]
True
```
```
>>> get_div3or4(list(range(100,150,2)))==[100, 102, 104, 108, 112, 114, 116, 120, 124, 126, 128, 132, 136, 138, 140, 144, 148]
True
```
```
>>> get_div3or4([1,2])==[]
True
```


Tests for q4 (points: [1, 1, 1, 1])
```
>>> my_median([0,1,2,3,4])==2
True
```
```
>>> my_median([1.0,1.1,2.0,3.0,3.9,4.1])==2.5
True
```
```
>>> inputList= [5, 2, 6.9, 3, 4, 0.1]
>>> my_median(inputList)==3.5
True
```
```
>>> inputList= [1.8, 'nine', 2, None, 2,2.7, 1+3J, 3.1, 4.3, 5]
>>> my_median(inputList)==2.7
True
```


Tests for q5 (points: 5)
```
>>> hex2dec("FFFFFF")==(2**24-1)
True
```
```
>>> hex2dec('ABCDEF')== 11259375
True
```
```
>>> hex2dec('12')==18
True
```
```
>>> hex2dec('ff')==255
True
```
```
>>> hex2dec('A0')==160
True
```


Tests for q6 (points: [1, 1, 1, 1, 1, 1])
```
>>> is_prime(3.3)==None
3.3 is not an integer greater than 1.
True
```
```
>>> is_prime(17)==True
True
```
```
>>> is_prime(1)==None
1 is not an integer greater than 1.
True
```
```
>>> is_prime(1001)==False
True
```


Tests for q7 (points: [0.5, 0.5, 0.5, 0.5, 1])
```
>>> len(random_sample)==40
True
```
```
>>> all((random_sample<=100+4*10) & (random_sample>=100-4*10))
True
```
```
>>> type(is_ge_mean[0])==np.bool_
True
```
```
>>> len(is_ge_mean)==40
True
```
```
>>> all(low_sample==random_sample[random_sample<100])
True
```


Tests for q8 (points: 3)
```
>>> sample_mean == np.sum(random_sample)/len(random_sample) and sample_mean==random_sample.mean()
True
```
```
>>> sample_var == np.var(random_sample,ddof=1)
True
```
```
>>> sample_sd == np.std(random_sample, ddof=1)
True
```


Tests for q9 (points: 5)
```
>>> np.all(even_matrix==np.array([[ 2,  4,  6,  8, 10],
...                               [12, 14, 16, 18, 20],
...                               [22, 24, 26, 28, 30],
...                               [32, 34, 36, 38, 40],
...                               [42, 44, 46, 48, 50],
...                               [52, 54, 56, 58, 60]]))
True
```
```
>>> np.all(odd_matrix==np.array([[ 1,  3,  5,  7,  9],
...                              [11, 13, 15, 17, 19],
...                              [21, 23, 25, 27, 29],
...                              [31, 33, 35, 37, 39],
...                              [41, 43, 45, 47, 49],
...                              [51, 53, 55, 57, 59]]))
True
```
```
>>> np.all(diff_mat==np.array([[1, 1, 1, 1, 1],
...                            [1, 1, 1, 1, 1],
...                            [1, 1, 1, 1, 1],
...                            [1, 1, 1, 1, 1],
...                            [1, 1, 1, 1, 1],
...                            [1, 1, 1, 1, 1]]))
True
```
```
>>> np.all(combined_by_row==np.array([[ 1,  3,  5,  7,  9],
...                                   [11, 13, 15, 17, 19],
...                                   [21, 23, 25, 27, 29],
...                                   [31, 33, 35, 37, 39],
...                                   [41, 43, 45, 47, 49],
...                                   [51, 53, 55, 57, 59],
...                                   [ 2,  4,  6,  8, 10],
...                                   [12, 14, 16, 18, 20],
...                                   [22, 24, 26, 28, 30],
...                                   [32, 34, 36, 38, 40],
...                                   [42, 44, 46, 48, 50],
...                                   [52, 54, 56, 58, 60]]))
True
```
```
>>> np.all(combined_by_col==np.array([[ 1,  3,  5,  7,  9,  2,  4,  6,  8, 10],
...                                   [11, 13, 15, 17, 19, 12, 14, 16, 18, 20],
...                                   [21, 23, 25, 27, 29, 22, 24, 26, 28, 30],
...                                   [31, 33, 35, 37, 39, 32, 34, 36, 38, 40],
...                                   [41, 43, 45, 47, 49, 42, 44, 46, 48, 50],
...                                   [51, 53, 55, 57, 59, 52, 54, 56, 58, 60]]))
True
```
